**We're going to help the guards out and compute whether a collection of swallows can carry a collection of coconuts between them. We know that "a five ounce bird could not carry a one pound coconut," so we will assume that a 5.5 ounce bird can carry a one pound coconut. More specifically, we will assume that every 5.5 ounces of bird can carry one pound of coconut.**

Prompt the user for the number of ounces of bird 
by asking: "How many ounces of birds are carrying the coconuts? ".
Prompt the user for the number of pounds of coconut by asking: 
"How many pounds of coconuts are there? ".
Remember to convert these values to floats! 
You can assume that the user input is formatted correctly.

<!-- grade_id: cell-cf9726a45fb25ba2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# def carrying_coconuts(ounces_of_birds, pounds_of_coconuts):

### BEGIN SOLUTION

def carrying_coconuts():
    ounces_of_birds = float(input('How many ounces of birds are carrying the coconuts?'))
    pounds_of_coconuts = float(input('How many pounds of coconuts are there?'))

    if ounces_of_birds / pounds_of_coconuts >= 5.5:
         print("Yes! Carrying the coconuts is possible.")
    else: 
        print("No. Carrying the coconuts is impossible.")
    
### END SOLUTION

carrying_coconuts()
```

Output:
```
How many ounces of birds are carrying the coconuts? 6
How many pounds of coconuts are there? 4
No. Carrying the coconuts is impossible.
```

# **Sample run**

How many ounces of birds are carrying the coconuts? 5
How many pounds of coconuts are there? 1
No. Carrying the coconuts is impossible.


How many ounces of birds are carrying the coconuts? 6.2
How many pounds of coconuts are there? 1.1
Yes! Carrying the coconuts is possie

py
How many ounces of birds are carrying the coconuts? 17
How many pounds of coconuts are there? 3
Yes! Carrying the cos is conutpossi
ts.py
How many ounces of birds are carrying the coconuts? 12.5
How many pounds of coconuts are there? 2.5
No. Carrying the coconuts is impossible.


<!-- grade_id: cell-308136d4df2f8427, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

import mock

def testing_coconuts():
    with mock.patch('builtins.input', side_effect=['5.5', '1']):
        assert carrying_coconuts() == 'Yes! Carrying the coconuts is possible.'
        
    with mock.patch('builtins.input', side_effect=['46', '25']):
        assert carrying_coconuts() == 'No. Carrying the coconuts is impossible.'
        
    with mock.patch('builtins.input', side_effect=['0', '1']):
        assert carrying_coconuts() == 'No. Carrying the coconuts is impossible.'
        
    with mock.patch('builtins.input', side_effect=['1000', '25']):
        assert carrying_coconuts() == 'Yes! Carrying the coconuts is possible.'
        
testing_coconuts()

### END HIDDEN TESTS
```

Output:
```
Yes! Carrying the coconuts is possible.
Cell In[26], line 18
     15     with mock.patch('builtins.input', side_effect=['1000', '25']):
     16         assert carrying_coconuts() == 'Yes! Carrying the coconuts is possible.'
---> 18 testing_coconuts()

Cell In[26], line 7, in testing_coconuts()
      5 def testing_coconuts():
      6     with mock.patch('builtins.input', side_effect=['5.5', '1']):
----> 7         assert carrying_coconuts() == 'Yes! Carrying the coconuts is possible.'
      9     with mock.patch('builtins.input', side_effect=['46', '25']):
     10         assert carrying_coconuts() == 'No. Carrying the coconuts is impossible.'

AssertionError:
```

Make a list of cheeses containing "Muenster", "Cheddar", and "Red Leicester"; these are the cheeses you have in your shop.

Begin by printing the string "Good morning. Welcome to the National Cheese Emporium!" to the console.

Next, in a loop, we will repeatedly ask the user which cheese they would like to buy. If the user enters the exact name of a cheese that you have, affirm that you have the cheese they asked for in the format "We have {}, yessir.". If the user does not enter the exact name of a cheese that you have, say that you don't have that cheese in the form "I'm afraid we don't have any {}." and then reprompt the user to ask for another cheese.

The user is also allowed to enter special questions. If the user enters either of the strings "You... do have some cheese, don't you?" or "Have you in fact got any cheese here at all?", you must reply by listing the number of cheeses that you have in the format "We have {} cheese(s)!", along with the name of each type of cheese you have, one on each

**For anything that is not specified in the above specification, your program can behave in anyway you'd like. For example, you can customize the prompts and messages.** line.

<!-- grade_id: cell-274ac8afe3538dae, points: 0, grade: false, solution: true -->
```python
# YOUR SOLUTION HERE

### BEGIN SOLUTION

def chees():
    
    cheeses = ["Muenster", "Cheddar", "Red Leicester"]
    print("Good morning. Welcome to the National Cheese Emporium!")

    while True:
        user_input = input("Which cheese would you like to buy? ")

        if user_input == "You... do have some cheese, don't you?" or user_input == "Have you in fact got any cheese here at all?":
            print(f'We have {len(cheeses)} cheese(s)!')
            for chees in cheeses: 
                print(chees)
            break

        elif user_input in cheeses:
            print("We have {}, yessir.".format(user_input))
            break

        else:
            print("I'm afraid we don't have any {}.".format(user_input))
### END SOLUTION

```

**Sample Run**

Good morning. Welcome to the National Cheese Emporium!
What would you like? Red Windsor
I'm afraid we don't have any Red Windsor.
What would you like? Lancashire
I'm afraid we don't have any Lancashire.
What would you like? cheddar
I'm afraid we don't have any cheddar.
What would you like? Cheddar
We have Cheddar, yessir.

Good morning. Welcome to the National Cheese Emporium!
What would you like? Red Windsor
I'm afraid we don't have any Red Windsor.
What would you like? cHeDdAr
I'm afraid we don't have any cHeDdAr.
What would you like? Have you in fact got any cheese here at all?
We have 3 cheese(s)!
Muenster
Cheddar
Red Leicester
What would you like? exit
I'm afraid we don't have any exit.
What would you like? LET ME OUT
I'm afraid we don't have any LET ME OUT.
What would you like? Cheddar
We have Cheddar, yessir.

<!-- grade_id: cell-5057de5b041f6f13, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest
import io
import mock

def testing_cheese():
    with mock.patch('builtins.input', return_value='Cheddar'):
        assert chees() == 'We have Cheddar, yessir.'
        
    with mock.patch('builtins.input', return_value='Muenster'):
        assert chees() == 'We have Muenster, yessir.'
        
    with mock.patch('builtins.input', return_value='Red Leicester'):
        assert chees() == 'We have Red Leicester, yessir.'

    with mock.patch('builtins.input', return_value='gift'):
        assert chees() == "I'm afraid we don't have any gift."
        
    with mock.patch('builtins.input', return_value='Chedder'):
        assert chees() == "I'm afraid we don't have any Chedder."

    # tc = unittest.TestCase()
    # with mock.patch('sys.stdout', new_callable=io.StringIO) as mock_stdout:
      #  tc.AssertEqual(mock_stdout.get_value(), 'We have 3 cheese(s)!\nMuenster\nCheddar\nRed Leicester')
        
    with mock.patch('builtins.input', return_value="You... do have some cheese, don't you?"):
        assert chees() == "We have 3 cheese(s)!\nMuenster\nCheddar\nRed Leicester"

testing_cheese()
### END HIDDEN TESTS
```

Output:
```
Good morning. Welcome to the National Cheese Emporium!
We have Cheddar, yessir.
Cell In[12], line 29
     26     with mock.patch('builtins.input', return_value="You... do have some cheese, don't you?"):
     27         assert chees() == "We have 3 cheese(s)!\nMuenster\nCheddar\nRed Leicester"
---> 29 testing_cheese()

Cell In[12], line 8, in testing_cheese()
      6 def testing_cheese():
      7     with mock.patch('builtins.input', return_value='Cheddar'):
----> 8         assert chees() == 'We have Cheddar, yessir.'
     10     with mock.patch('builtins.input', return_value='Muenster'):
     11         assert chees() == 'We have Muenster, yessir.'

AssertionError:
```

```python

```

```python

```

```python

```

```python

```