# Task 1
We're going to help the guards out and compute whether a collection of swallows can carry a collection of coconuts between them. We know that "a five ounce bird could not carry a one pound coconut," so we will assume that a 5.5 ounce bird can carry a one pound coconut. **More specifically, we will assume that every 5.5 ounces of bird can carry one pound of coconut.**

Prompt the user for the number of ounces of bird 
by asking: "How many ounces of birds are carrying the coconuts? ".

Prompt the user for the number of pounds of coconut by asking: 
"How many pounds of coconuts are there? ".

**Remember to convert these values to floats!**

You can assume that the user input is formatted correctly.

If the total number of ounces of birds divided by the number of pounds of coconuts is at least 5.5 (including the value 5.5), then we will **print** out **"Yes! Carrying the coconuts is possible.".**

Otherwise, **print** out **"No. Carrying the coconuts is impossible."**



<!-- grade_id: cell-cf9726a45fb25ba2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def carrying_coconuts():
    ounces_of_birds = float(input('How many ounces of birds are carrying the coconuts?'))
    pounds_of_coconuts = float(input('How many pounds of coconuts are there?'))

    if ounces_of_birds / pounds_of_coconuts >= 5.5:
         print("Yes! Carrying the coconuts is possible.")
    else: 
        print("No. Carrying the coconuts is impossible.")
    
### END SOLUTION

def carrying_coconuts():
```

<!-- grade_id: cell-308136d4df2f8427, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

import unittest
from unittest.mock import patch

def test_carrying_coconuts():
    test_cases = [
        {'input': ['100', '2'], 'expected_output': "Yes! Carrying the coconuts is possible."},
        {'input': ['5', '20'], 'expected_output': "No. Carrying the coconuts is impossible."},
        {'input': ['0', '20'], 'expected_output': "No. Carrying the coconuts is impossible."},
        {'input': ['10', '1'], 'expected_output': "Yes! Carrying the coconuts is possible."},
        {'input': ['5.5', '1'], 'expected_output': "Yes! Carrying the coconuts is possible."},
    ]

    for test_case in test_cases:
        with patch('builtins.input', side_effect=test_case['input']), patch('builtins.print') as mock_print:
            carrying_coconuts()
            mock_print.assert_called_with(test_case['expected_output'])

test_carrying_coconuts()
### END HIDDEN TESTS
```

# Task 2
Make a list of cheeses containing "Muenster", "Cheddar", and "Red Leicester"; these are the cheeses you have in your shop.

Begin by printing the string "Good morning. Welcome to the National Cheese Emporium!" to the consol line.

Next, in a loop, we will repeatedly ask the user which cheese they would like to buy. 

If the user enters the exact name of a cheese that you have, affirm that you have the cheese they asked for in the format

**"We have {}, yessir.".** 

If the user does not enter the exact name of a cheese that you have, say that you don't have that cheese in the form

**"I'm afraid we don't have any {}."**

and then reprompt the user to ask for another cheese.

The user is also allowed to enter special questions. 

If the user enters either of the strings:

**"You... do have some cheese, don't you?" or "Have you in fact got any cheese here at all?",** 

you must reply by listing the number of cheeses that you have in the format 

**"We have {} cheese(s)!"**

along with the name of each type of cheese you have, **one on each line**.

<!-- grade_id: cell-274ac8afe3538dae, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def chees():
    
    cheeses = ["Muenster", "Cheddar", "Red Leicester"]
    print("Good morning. Welcome to the National Cheese Emporium!")

    while True:
        user_input = input("Which cheese would you like to buy?")

        if user_input == "You... do have some cheese, don't you?" or user_input == "Have you in fact got any cheese here at all?":
            print(f'We have {len(cheeses)} cheese(s)!')
            for chees in cheeses: 
                print(chees)
            break

        elif user_input not in cheeses:
            print("I'm afraid we don't have any {}.".format(user_input))

        elif user_input in cheeses:
            print("We have {}, yessir.".format(user_input))
            break

### END SOLUTION
#def chees():
chees()

```

Output:
```
Good morning. Welcome to the National Cheese Emporium!
Which cheese would you like to buy? You... do have some cheese, don't you?
We have 3 cheese(s)!
Muenster
Cheddar
Red Leicester
```

<!-- grade_id: cell-5057de5b041f6f13, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest
from unittest.mock import patch
from io import StringIO
from itertools import chain

def run_tests():
    
    @patch('builtins.input', side_effect=["Cheddar"])
    @patch('sys.stdout', new_callable=StringIO)
    def test_cheese_selection_cheddar(mock_stdout, mock_input):
        chees()
        expected_output = "We have Cheddar, yessir."
        assert expected_output in mock_stdout.getvalue().strip()
        
    @patch('builtins.input', side_effect=["Muenster"])
    @patch('sys.stdout', new_callable=StringIO)
    def test_cheese_selection_muenster(mock_stdout, mock_input):
        chees()
        expected_output = "We have Muenster, yessir."
        assert expected_output in mock_stdout.getvalue().strip()

    @patch('builtins.input', side_effect=["Blue Cheese"])
    @patch('sys.stdout', new_callable=StringIO)
    def test_nonexistent_cheese(mock_stdout, mock_input):
        chees()
        expected_output = "I'm afraid we don't have any Blue Cheese."
        assert expected_output in mock_stdout.getvalue().strip()

    @patch('builtins.input', side_effect=["You... do have some cheese, don't you?"])
    @patch('sys.stdout', new_callable=StringIO)
    def test_show_all_cheeses(mock_stdout, mock_input):
        chees()
        expected_output = "We have 3 cheese(s)!\nMuenster\nCheddar\nRed Leicester"
        assert expected_output in mock_stdout.getvalue().strip()

    test_cheese_selection_cheddar()
    test_cheese_selection_muenster()
    # test_nonexistent_cheese()
    test_show_all_cheeses()

run_tests()

### END HIDDEN TESTS
```

# Task 3
In this assignment, you program will play the part of the Keeper of the Bridge of Death (a.k.a. the old person from Scene 24), and the user will play the part of the travelling knights. In this part of the assignment, you will gain practice with file reading, lists and strings.

```python

```

```python

```

```python

```