# Task 1
Edit say_hello so that it **prints** "Hello, world!" when executed.



<!-- grade_id: cell-c4a5f4669af2b450, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def say_hello():
     print('Hello, world!')
### END SOLUTION
def say_hello():

```

<!-- grade_id: cell-1971fb7d65d88972, points: 0, grade: true, solution: false -->
```python
# Run this cell to check your code

import unittest
import unittest.mock
import io

tc = unittest.TestCase()
with unittest.mock.patch('sys.stdout', new_callable=io.StringIO) as mock_stdout:
    say_hello()
    tc.assertNotEqual(mock_stdout.getvalue(), 'Hello, world\n')

### BEGIN HIDDEN TESTS

tc = unittest.TestCase()
with unittest.mock.patch('sys.stdout', new_callable=io.StringIO) as mock_stdout:
    say_hello()
    tc.assertEqual(mock_stdout.getvalue(), 'Hello, world!\n')

### END HIDDEN TESTS
```

# Task 2
**Edit fizzbuzz(n)** to find the **sum of all the multiples of 3 or 5 below 1001.** As a sanity check, the last two digits of the sum should be 68.

Make sure that your function **returns** the sum.

<!-- grade_id: cell-8ed2007c3a947eca, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def fizzbuzz(number):
    return sum([i for i in range(number) if (i % 3 == 0) or (i % 5 == 0)])
    
### END SOLUTION
def fizzbuzz(n):
```

Output:
```
Cell In[3], line 7
    def fizzbuzz(n):
                    ^
SyntaxError: incomplete input
```

<!-- grade_id: cell-d514062f25ca23b9, points: 0, grade: true, solution: false -->
```python
### Run this cell to check your function with other values
assert fizzbuzz(100) == 2318
assert fizzbuzz(56) == 753
assert fizzbuzz(1) == 0
assert fizzbuzz(2356) == 1295250

### BEGIN HIDDEN TESTS

def testing_fizzbuzz():
    assert fizzbuzz(1001) == 234168

testing_fizzbuzz()

### END HIDDEN TESTS
```

# Task 3
The Collatz sequence is an iterative sequence defined on the positive integers by:

n -> n / 2 if n is **even**

n -> 3n + 1 if n is **od**d

For example, using the rule above and starting with 13 yields the sequence:

**13 -> 40 -> 20 -> 10 -> 5 -> 16 -> 8 -> 4 -> 2 -> 1**

It can be seen that this sequence (starting at 13 and finishing at 1) contains 10 terms. Although unproven, it it hypothesized that all starting numbers finish at 1.

**Edit max_collatz_len(n)** so that **it computes the the length of the longest chain** which has a starting number under n and use it to find the longest chain that has a starting number under 1000.

**Note** it is easier to implement a **collatz_len()** function to calculate the collatz length and then to edit a **max_collatz_len()** function so that you can finf the longest chain.

<!-- grade_id: cell-9eb47e90c445335d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def collatz_len(n):
    count = 1
    while n != 1:
        n = n // 2 if n % 2 == 0 else 3*n+1
        count += 1
    return count

def max_collatz_len(n):
    return max(collatz_len(x) for x in range(1, n))

### END SOLUTION
def collatz_len(n):
def max_collatz_len(n):
```

<!-- grade_id: cell-c9d4008917832eec, points: 0, grade: true, solution: false -->
```python
### Run this cell to check your function with other values
assert max_collatz_len(100) == 119
assert max_collatz_len(2) == 1
assert max_collatz_len(50) == 112
### BEGIN HIDDEN TESTS

def test_max_collatz_len():
    assert max_collatz_len(10) == 20
    assert max_collatz_len(20) == 21
    assert max_collatz_len(1000) == 179
    
test_max_collatz_len()

### END HIDDEN TESTS
```

For the next question you just need to re-write **collatz_len()** function.

<!-- grade_id: cell-1d89288ce2c94977, points: 0, grade: false, solution: true -->
```python
'''You can also try to solve a challenge problem'''

'''Challenge: Same question, but for any starting number under 1,000,000. What about for any starting number under 10,000,000? 
You may need to implement a cleverer-than-naive algorithm.'''

### BEGIN SOLUTION
from functools import lru_cache


@lru_cache(maxsize=None)
def collatz_len(n):
    if n == 1:
        return 1
    else:
        n_next = 3*n+1 if n % 2 == 1 else n // 2
        return collatz_len(n_next) + 1
        
### END SOLUTION
def collatz_len(n):
```

<!-- grade_id: cell-dcda42be96acc037, points: 0, grade: true, solution: false -->
```python
'''Run this cell only if you are sure about the solution or it will take a huge amount of time to compute.'''
assert max_collatz_len(1000000) == 525
assert max_collatz_len(10000000) == 686
```

# Task 4
**Implement print_tictactoe using print() so that, when run, it prints out a tic-tac-toe board.**


```python
''' X | . | . 
   -----------
    . | O | . 
   -----------
    . | O | X ''' 
```

<!-- grade_id: cell-088d1d857b8aa75a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def print_tic_tac_toe():
    board = [[' X', '.', '. '], [' .', 'O', '. '], [' .', 'O', 'X']]
    for i, row in enumerate(board):
        print(" | ".join(row))
        if i < len(board) - 1:
            print("-" * 11)
### END SOLUTION
def print_tictactoe():
```

<!-- grade_id: cell-f7dbf9185beec70b, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from io import StringIO
import sys

def capture_print_output(func):
    captured_output = StringIO()
    sys.stdout = captured_output
    func()
    sys.stdout = sys.__stdout__
    return captured_output.getvalue().strip()

def test_print_tic_tac_toe():
    expected_output = "X | . | . \n-----------\n . | O | . \n-----------\n . | O | X"
    actual_output = capture_print_output(print_tic_tac_toe)
    assert actual_output == expected_output, "Printed output does not match expected output."

test_print_tic_tac_toe()

### END HIDDEN TESTS
```

# Task 5
**Write a program that prints out a SUPER tic-tac-toe board.**

```python
'''  |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   ========+========+========
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   ========+========+========
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--  
     |  |  H  |  |  H  |  |   '''
```

<!-- grade_id: cell-e215a53497362d23, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def super_tic_tac_toe():
    print('  |  |  H  |  |  H  |  |  \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
========+========+========  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
========+========+========  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |  ')

### END SOLUTION
# def super_tic_tac_toe():
super_tic_tac_toe()
```

<!-- grade_id: cell-0f30cc9cdc88bdfa, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS 
def test_board_output():
    expected_output = (
        '  |  |  H  |  |  H  |  |  \n\
         --+--+--H--+--+--H--+--+--  \n\
           |  |  H  |  |  H  |  |    \n\
         --+--+--H--+--+--H--+--+--  \n\
           |  |  H  |  |  H  |  |    \n\
         ========+========+========  \n\
           |  |  H  |  |  H  |  |    \n\
         --+--+--H--+--+--H--+--+--  \n\
           |  |  H  |  |  H  |  |    \n\
         --+--+--H--+--+--H--+--+--  \n\
           |  |  H  |  |  H  |  |    \n\
         ========+========+========  \n\
           |  |  H  |  |  H  |  |    \n\
         --+--+--H--+--+--H--+--+--  \n\
           |  |  H  |  |  H  |  |    \n\
         --+--+--H--+--+--H--+--+--  \n\
           |  |  H  |  |  H  |  |  ')
    assert super_tic_tac_toe() == expected_output

### END HIDDEN TESTS
test_board_output()
```

# Task 6
This problem is called lattice paths. Starting in the top left corner of a 2*2 grid, and only being able to move to the right and down, there are exactly 6 routes to the bottom right corner.

Implement **lattice_paths(m, n)** to be a recursive function that, when run, calculates the number of paths from the top left of an m*n grid (m rows and n columns) to the bottom right corner.

<!-- grade_id: cell-ad3f4f6f3038d5d2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def lattice_paths(m, n):
    # Base case
    if m == 0 or n == 0:
        return 1
    # Recursive case
    else:
        return lattice_paths(m - 1, n) + lattice_paths(m, n - 1)
        
### END SOLUTION
def lattice_paths(m, n):
```

<!-- grade_id: cell-61ad69aff9c9fb5e, points: 0, grade: true, solution: false -->
```python
# Run this cell to check how your function works
assert lattice_paths(6, 6) == 924
assert lattice_paths(7, 7) == 3432

### BEGIN HIDDEN TESTS
assert lattice_paths(1, 1) == 2
assert lattice_paths(0, 0) == 1
assert lattice_paths(10, 10) == 184756
### END HIDDEN TESTS
```

Notice that the way that our code is constructed results in a lot of repeated calls to the same functions. For example, if you start at the top of a 20*20 grid and go down, then right, you'll make a function call for a 19*19 grid. But you'll also make that call if you go right and then down. This is where caching comes in. We'll store the output of lattice_paths in a dictionary the first time that we call it and then the second time, we'll just immediately return the value from the dictionary so that we won't make calls twice.

In **lattice_paths_cached**, let's cache lattice_paths by adding its output to a dictionary. This dictionary should have keys that are tuples (m, n) and values that are the output lattice_paths(m,n). Make sure to check if the tuple (m,n) is in the cache dictionary at the beginning of the function execution.

<!-- grade_id: cell-524dfd01424c69af, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def lattice_paths_cached(m, n, memo={}):
    # Check if the result is already memoized
    if (m, n) in memo:
        return memo[(m, n)]
    
    # Base case
    if m == 0 or n == 0:
        return 1
    
    # Recursive case
    paths = lattice_paths_cached(m - 1, n, memo) + lattice_paths_cached(m, n - 1, memo)
    
    # Memoize the result before returning
    memo[(m, n)] = paths
    
    return paths
### END SOLUTION
def lattice_paths_cached(m, n, memo={}):
```

<!-- grade_id: cell-aa464037852cfb77, points: 0, grade: true, solution: false -->
```python
'''Only run this cell if you are sure about your code or it will take a huge amount of time to compute'''

assert lattice_paths_cached(40, 40) == 107507208733336176461620
assert lattice_paths_cached(100, 100) == 90548514656103281165404177077484163874504589675413336841320
```

# Task 5
Write a function to compute the GCD of two positive integers. You can freely use the fact that gcd(a, b) is mathematically equal to gcd(b, a % b), and that gcd(a, 0) == a.

You can assume that a >= b if you'd like. But remember that there might be tests with negative numbers.

It is possible to accomplish this in three lines of Python code (or with extra cleverness, even fewer!). Consider exploiting tuple packing and unpacking!

**Implement your solution in the gcd() function.**

<!-- grade_id: cell-7f4f96e22d5ebe2f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def gcd(a, b):
    return a if b == 0 else gcd(b, a % b)

### END SOLUTION
def gcd(a, b):
```

<!-- grade_id: cell-28d093e70a2466c1, points: 0, grade: true, solution: false -->
```python
# Run this cell to check how your function works
assert gcd(48, 18) == 6
assert gcd(123456, 789012) == 12
assert gcd(17, 19) == 1

### BEGIN HIDDEN TESTS
assert gcd(48, 0) == 48
assert gcd(0, 0) == 0
assert gcd(-36, 48) == 12
### END HIDDEN TESTS

```

# Task 6

```python
'''I asked our course staff what their favorite animal is and they gave me these answers:'''
fav_animals = {
    'parth': 'unicorn',
    'michael': 'elephant',
    'sam': 'ox',
    'zheng': 'tree',
    'theo': 'puppy',
    'alex': 'dog',
    'nick': 'daisy'
}
```

```python
'''I know that, realistically, they're lying to themselves. In fact, the dict probably looks more like this:'''
fav_animals = {
    'parth': 'unicorn',
    'michael': 'unicorn',
    'sam': 'unicorn',
    'zheng': 'tree',
    'theo': 'unicorn',
    'alex': 'dog',
    'nick': 'daisy'
}
```

In this problem, we'll reverse the fav_animals dictionary to create a new dictionary which associates animals to a list of people for whom that animal is their favorite.

More precisely, **write a function that properly reverses the keys and values of a dictionary** - each key (originally a value) should map to a collection of values (originally keys) that mapped to it. Example in the cell:

```python
flip_dict({"CA": "US", "NY": "US", "ON": "CA"}) => {"US": ["CA", "NY"], "CA": ["ON"]}
```

<!-- grade_id: cell-423c5e21cf539125, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def flip_dict(d):
    out = {}
    for key, value in d.items():
        if value not in out:
            out[value] = []
        out[value].append(key)
    return out

### END SOLUTION
def flip_dict(d):
```

<!-- grade_id: cell-f7e4b164b1cee83b, points: 0, grade: true, solution: false -->
```python
# Run this cell to check how your function works
assert flip_dict({"CA": "US", "NY": "US", "ON": "CA"}) == {"US": ["CA", "NY"], "CA": ["ON"]}

### BEGIN HIDDEN TESTS

assert flip_dict({
    'parth': 'unicorn',
    'michael': 'unicorn',
    'sam': 'unicorn',
    'zheng': 'tree',
    'theo': 'unicorn',
    'alex': 'dog',
    'nick': 'daisy'}) == {'unicorn': ['parth', 'michael', 'sam', 'theo'],
                          'tree': ['zheng'],
                          'dog': ['alex'],
                          'daisy': ['nick']}

### END HIDDEN TESTS
```

# Bonus problems #

# Bonus task 1 
Write a function with the name generate_pascal_row that generates the next level of Pascal's triangle given a list that represents a row of Pascal’s triangle.

generate_pascal_row([1, 2, 1]) -> [1, 3, 3, 1]

generate_pascal_row([1, 4, 6, 4, 1]) -> [1, 5, 10, 10, 5, 1]


generate_pascal_row([]) ->[
1]

As a reminder, each element in a row of Pascal's triangle is formed by summing the two elements in the previous row directly above (to the left and right) that elements. If there is only one element directly above, we only add that one. For example, the first 5 rows of Pascal's triangle look like:

```python
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

<!-- grade_id: cell-2d512bb98118c6d3, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def generate_pascal_row(row):
    if not row:
        return [1] 
    
    next_row = [1] + [sum(pair) for pair in zip(row, row[1:])] + [1]
    
    return next_row
    
### END SOLUTION
def generate_pascal_row(row):
```

<!-- grade_id: cell-7783c6b713759950, points: 0, grade: true, solution: false -->
```python
assert generate_pascal_row([1, 2, 1]) == [1, 3, 3, 1]
assert generate_pascal_row([1, 4, 6, 4, 1]) == [1, 5, 10, 10, 5, 1]

### BEGIN HIDDEN TESTS 
assert generate_pascal_row([]) == [1]
assert generate_pascal_row([1]) == [1, 1]
assert generate_pascal_row([1, 1]) == [1, 2, 1]
### END HIDDEN TESTS

```

# Bonus task 2

**Implement pretty_pascal_triangle so that given a number n, it prints out the first n rows of Pascal's triangle, centering each line.** You should use the generate_pascal_row function you wrote previously. The Pascal's triangle with 1 row just contains the number 1.

To center a string in Python, you can use use string format specifiers. '{:^10}'.format(var) will produce a string of length 10 or len(var) (whichever is longer), with str(var) centered.

```python
'{:^10}'.format('CS41') # => '   CS41   '
```

<!-- grade_id: cell-68777f9b3f94a8a4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def pretty_pascal_triangle(n):
    prev = []
    rows = [prev := generate_pascal_row(prev) for _ in range(n)]
    
    rows = [[str(e) for e in row] for row in rows]
    rows = [' '.join(row) for row in rows]
    width = len(rows[-1])
    
    # Print each row
    for row in rows:
        print("{row:^{width}}".format(row=row, width=width))

pretty_pascal_triangle(10)
### END SOLUTION
#def pretty_pascal_triangle(n):
```

<!-- grade_id: cell-c2c7b32563f0f270, points: 0, grade: true, solution: false -->
```python
# Run this cell to check how your function works
tc = unittest.TestCase()
with unittest.mock.patch('sys.stdout', new_callable=io.StringIO) as mock_stdout:
    pretty_pascal_triangle(1)
    tc.assertEqual(mock_stdout.getvalue(), "1\n")

tc1 = unittest.TestCase()
with unittest.mock.patch('sys.stdout', new_callable=io.StringIO) as mock_stdout:
    pretty_pascal_triangle(3)
    tc1.assertEqual(mock_stdout.getvalue(), "  1  \n 1 1 \n1 2 1\n")

             
### BEGIN HIDDEN TESTS
tc3 = unittest.TestCase()
with unittest.mock.patch('sys.stdout', new_callable=io.StringIO) as mock_stdout:
    pretty_pascal_triangle(5)
    tc3.assertEqual(mock_stdout.getvalue(), "    1    \n   1 1   \n  1 2 1  \n 1 3 3 1 \n1 4 6 4 1\n")
### END HIDDEN TESTS
```

```python
print("             1             \n            1 1            \n           1 2 1           \n          1 3 3 1          \n         1 4 6 4 1         \n       1 5 10 10 5 1       \n     1 6 15 20 15 6 1     \n    1 7 21 35 35 21 7 1    \n  1 8 28 56 70 56 28 8 1  \n1 9 36 84 126 126 84 36 9 1\n")
```

```python
'''Run this cell to have a list of words at your working directory'''

import requests

file_url = 'https://www.mit.edu/~ecprice/wordlist.10000'

response = requests.get(file_url)

if response.status_code == 200:
    with open('words.txt', 'wb') as file:
        file.write(response.content)
    print('File downloaded successfully.')
else:
    print(f'Failed to download the file. Status code: {response.status_code}')
```

# Bonus task 3
**Write a function called load_english** that loads the words from this file and returns them. What would be an appropriate data structure in which to store these words?

<!-- grade_id: cell-79e0b2930b369945, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def load_english(file_path):
    try:
        with open(file_path, 'r', encoding='utf-8') as file:
            words = file.read().split()
        return words
    except FileNotFoundError:
        print(f"Error: File '{file_path}' not found.")
        return []
# END SOLUTION
def load_english(file_path):
```

<!-- grade_id: cell-a18880f7f017f3ad, points: 0, grade: true, solution: false -->
```python
# Run this cell to check if you are right with your solution
assert len(load_english('words.txt')) == 10000
```

```python

```

```python

```

```python

```