Edit say_hello so that it **prints** "Hello, world!" when executed.



<!-- grade_id: cell-c4a5f4669af2b450, points: 0, grade: false, solution: true -->
```python
### WRITE YOR CODE HERE
def say_hello():

### BEGIN SOLUTION
def say_hello():
     print('Hello, world!')
### END SOLUTION
```

Output:
```
Cell In[3], line 3
    raise NotImplementedError
    ^
IndentationError: expected an indented block after function definition on line 2
```

<!-- grade_id: cell-1971fb7d65d88972, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import unittest
import unittest.mock
import io

tc = unittest.TestCase()
with unittest.mock.patch('sys.stdout', new_callable=io.StringIO) as mock_stdout:
    say_hello()
    tc.assertEqual(mock_stdout.getvalue(), 'Hello, world!\n')

### END HIDDEN TESTS
```

**Edit fizzbuzz(n) to find the sum of all the multiples of 3 or 5 below 1001. As a sanity check, the last two digits of the sum should be 68.**

Make sure that your function **returns** the sum.

<!-- grade_id: cell-8ed2007c3a947eca, points: 0, grade: false, solution: true -->
```python
def fizzbuzz(n):

### BEGIN SOLUTION

def fizzbuzz(number):
    return sum([i for i in range(number) if (i % 3 == 0) or (i % 5 == 0)])
    
### END SOLUTION
```

Output:
```
1295250
```

<!-- grade_id: cell-d514062f25ca23b9, points: 0, grade: true, solution: false -->
```python
### Run this cell to check your function with other values
assert fizzbuzz(100) == 2318
assert fizzbuzz(56) == 753
assert fizzbuzz(1) == 0
assert fizzbuzz(2356) == 1295250


### BEGIN HIDDEN TESTS

def testing_fizzbuzz():
    assert fizzbuzz(1001) == 234168

testing_fizzbuzz()

### END HIDDEN TESTS
```

**Note** it is easier to implement a function to calculate the collatz length and then to edit a max_collatz_len function

<!-- grade_id: cell-9eb47e90c445335d, points: 0, grade: false, solution: true -->
```python
def max_collatz_len(n):

### BEGIN SOLUTION

def collatz_len(n):
    count = 1
    while n != 1:
        n = n // 2 if n % 2 == 0 else 3*n+1
        count += 1
    return count

def max_collatz_len(n):
    return max(collatz_len(x) for x in range(1, n))

### END SOLUTION
```

Output:
```
112
```

<!-- grade_id: cell-c9d4008917832eec, points: 0, grade: true, solution: false -->
```python
### Run this cell to check your function with other values
assert max_collatz_len(100) == 119
assert max_collatz_len(2) == 1
assert max_collatz_len(50) == 112
### BEGIN HIDDEN TESTS

def test_max_collatz_len():
    assert max_collatz_len(10) == 20
    assert max_collatz_len(20) == 21
    assert max_collatz_len(1000) == 179
    
test_max_collatz_len()

### END HIDDEN TESTS
```

<!-- grade_id: cell-1d89288ce2c94977, points: 0, grade: false, solution: true -->
```python
'''You can also try to solve a challenge problem'''

'''Challenge: Same question, but for any starting number under 1,000,000. What about for any starting number under 10,000,000? 
You may need to implement a cleverer-than-naive algorithm.'''

### BEGIN SOLUTION
from functools import lru_cache


@lru_cache(maxsize=None)
def collatz_len(n):
    if n == 1:
        return 1
    else:
        n_next = 3*n+1 if n % 2 == 1 else n // 2
        return collatz_len(n_next) + 1
        
### END SOLUTION
```

<!-- grade_id: cell-dcda42be96acc037, points: 0, grade: true, solution: false -->
```python
'''Run this cell only if you are sure about the solution or it will take a huge amount of time to compute.'''
assert max_collatz_len(1000000) == 525
```

**Implement print_tictactoe using print() so that, when run, it prints out a tic-tac-toe board.**


```python
''' X | . | . 
   -----------
    . | O | . 
   -----------
    . | O | X ''' 
```

Output:
```
' X | . | . \n   -----------\n    . | O | . \n   -----------\n    . | O | X '
```

<!-- grade_id: cell-088d1d857b8aa75a, points: 0, grade: false, solution: true -->
```python
def print_tictactoe():

### BEGIN SOLUTION
def print_tic_tac_toe():
    board = [[' X', '.', '. '], [' .', 'O', '. '], [' .', 'O', 'X']]
    for i, row in enumerate(board):
        print(" | ".join(row))
        if i < len(board) - 1:
            print("-" * 11)
### END SOLUTION


```

Output:
```
X | . | . 
-----------
 . | O | . 
-----------
 . | O | X
```

<!-- grade_id: cell-f7dbf9185beec70b, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from io import StringIO
import sys

def capture_print_output(func):
    captured_output = StringIO()
    sys.stdout = captured_output
    func()
    sys.stdout = sys.__stdout__
    return captured_output.getvalue().strip()

def test_print_tic_tac_toe():
    expected_output = "X | . | . \n-----------\n . | O | . \n-----------\n . | O | X"
    actual_output = capture_print_output(print_tic_tac_toe)
    assert actual_output == expected_output, "Printed output does not match expected output."

test_print_tic_tac_toe()

### END HIDDEN TESTS
```

**Write a program that prints out a SUPER tic-tac-toe board.**

```python
'''  |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   ========+========+========
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |  
   ========+========+========
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+-- 
     |  |  H  |  |  H  |  |  
   --+--+--H--+--+--H--+--+--
     |  |  H  |  |  H  |  |   '''
```

Output:
```
'  |  |  H  |  |  H  |  |  \n   --+--+--H--+--+--H--+--+--\n     |  |  H  |  |  H  |  |  \n   --+--+--H--+--+--H--+--+--\n     |  |  H  |  |  H  |  |  \n   ========+========+========\n     |  |  H  |  |  H  |  |  \n   --+--+--H--+--+--H--+--+--\n     |  |  H  |  |  H  |  |  \n   --+--+--H--+--+--H--+--+--\n     |  |  H  |  |  H  |  |  \n   ========+========+========\n     |  |  H  |  |  H  |  |  \n   --+--+--H--+--+--H--+--+-- \n     |  |  H  |  |  H  |  |  \n   --+--+--H--+--+--H--+--+--\n     |  |  H  |  |  H  |  |   '
```

<!-- grade_id: cell-e215a53497362d23, points: 0, grade: false, solution: true -->
```python
def super_tic_tac_toe():

### BEGIN SOLUTION
def super_tic_tac_toe():
    print('  |  |  H  |  |  H  |  |  \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
========+========+========  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
========+========+========  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |    \n\
--+--+--H--+--+--H--+--+--  \n\
  |  |  H  |  |  H  |  |  ')

super_tic_tac_toe()
### END SOLUTION
```

Output:
```
|  |  H  |  |  H  |  |  
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
========+========+========  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
========+========+========  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |
```

<!-- grade_id: cell-0f30cc9cdc88bdfa, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS 
from io import StringIO
import sys

def capture_print_output(func):
    captured_output = StringIO()
    sys.stdout = captured_output
    func()
    sys.stdout = sys.__stdout__
    return captured_output.getvalue().strip()

def test_super_tic_tac_toe():
    expected_output = '''  |  |  H  |  |  H  |  |  
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
========+========+========  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
========+========+========  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |    
--+--+--H--+--+--H--+--+--  
  |  |  H  |  |  H  |  |  '''

    actual_output = capture_print_output(super_tic_tac_toe)
    assert actual_output == expected_output, "Printed output does not match expected output."

# Run the test
test_super_tic_tac_toe()


### END HIDDEN TESTS
```

Output:
```
Cell In[3], line 35
     32     assert actual_output == expected_output, "Printed output does not match expected output."
     34 # Run the test
---> 35 test_super_tic_tac_toe()

Cell In[3], line 32, in test_super_tic_tac_toe()
     13     expected_output = '''  |  |  H  |  |  H  |  |  
     14 --+--+--H--+--+--H--+--+--  
     15   |  |  H  |  |  H  |  |    
   (...)
     28 --+--+--H--+--+--H--+--+--  
     29   |  |  H  |  |  H  |  |  '''
     31     actual_output = capture_print_output(super_tic_tac_toe)
---> 32     assert actual_output == expected_output, "Printed output does not match expected output."

AssertionError: Printed output does not match expected output.
```

This problem is called lattice paths. Starting in the top left corner of a 2*2 grid, and only being able to move to the right and down, there are exactly 6 routes to the bottom right corner.

Implement **lattice_paths(m, n)** to be a recursive function that, when run, calculates the number of paths from the top left of an m*n grid (m rows and n columns) to the bottom right corner.

<!-- grade_id: cell-ad3f4f6f3038d5d2, points: 0, grade: false, solution: true -->
```python
def lattice_paths(m, n):

### BEGIN SOLUTION

def lattice_paths(m, n):
    # Base case
    if m == 0 or n == 0:
        return 1
    # Recursive case
    else:
        return lattice_paths(m - 1, n) + lattice_paths(m, n - 1)
        
### END SOLUTION
```

Output:
```
184756
```

<!-- grade_id: cell-61ad69aff9c9fb5e, points: 0, grade: true, solution: false -->
```python
assert lattice_paths(6, 6) == 924
assert lattice_paths(7, 7) == 3432

### BEGIN HIDDEN TESTS
assert lattice_paths(1, 1) == 2
assert lattice_paths(0, 0) == 1
assert lattice_paths(10, 10) == 184756
### END HIDDEN TESTS
```

Notice that the way that our code is constructed results in a lot of repeated calls to the same functions. For example, if you start at the top of a 20*20 grid and go down, then right, you'll make a function call for a 19*19 grid. But you'll also make that call if you go right and then down. This is where caching comes in. We'll store the output of lattice_paths in a dictionary the first time that we call it and then the second time, we'll just immediately return the value from the dictionary so that we won't make calls twice.

In **lattice_paths_cached**, let's cache lattice_paths by adding its output to a dictionary. This dictionary should have keys that are tuples (m, n) and values that are the output lattice_paths(m,n). Make sure to check if the tuple (m,n) is in the cache dictionary at the beginning of the function execution.

<!-- grade_id: cell-524dfd01424c69af, points: 0, grade: false, solution: true -->
```python
def lattice_paths_cached(m, n, memo={}):

### BEGIN SOLUTION

def lattice_paths_cached(m, n, memo={}):
    # Check if the result is already memoized
    if (m, n) in memo:
        return memo[(m, n)]
    
    # Base case
    if m == 0 or n == 0:
        return 1
    
    # Recursive case
    paths = lattice_paths_cached(m - 1, n, memo) + lattice_paths_cached(m, n - 1, memo)
    
    # Memoize the result before returning
    memo[(m, n)] = paths
    
    return paths
### END SOLUTION
```

Output:
```
107507208733336176461620
```

<!-- grade_id: cell-aa464037852cfb77, points: 0, grade: true, solution: false -->
```python
'''Only run this cell if you are sure about your code or it will take a huge amount of time to compute'''

assert lattice_paths_cached(40, 40) == 107507208733336176461620
assert lattice_paths_cached(100, 100) == 90548514656103281165404177077484163874504589675413336841320
```

Write a function to compute the GCD of two positive integers. You can freely use the fact that gcd(a, b) is mathematically equal to gcd(b, a % b), and that gcd(a, 0) == a.

You can assume that a >= b if you'd like. But remember that there might be tests with negative numbers.

It is possible to accomplish this in three lines of Python code (or with extra cleverness, even fewer!). Consider exploiting tuple packing and unpacking!

**Implement your solution in the gcd() function.**

<!-- grade_id: cell-7f4f96e22d5ebe2f, points: 0, grade: false, solution: true -->
```python
def gcd(a, b):

### BEGIN SOLUTION

def gcd(a, b):
    return a if b == 0 else gcd(b, a % b)

### END SOLUTION

```

Output:
```
6
```

<!-- grade_id: cell-28d093e70a2466c1, points: 0, grade: true, solution: false -->
```python
assert gcd(48, 18) == 6
assert gcd(123456, 789012) == 12
assert gcd(17, 19) == 1

### BEGIN HIDDEN TESTS
assert gcd(48, 0) == 48
assert gcd(0, 0) == 0
assert gcd(-36, 48) == 12
### END HIDDEN TESTS

```

```python
'''I asked our course staff what their favorite animal is and they gave me these answers:'''
fav_animals = {
    'parth': 'unicorn',
    'michael': 'elephant',
    'sam': 'ox',
    'zheng': 'tree',
    'theo': 'puppy',
    'alex': 'dog',
    'nick': 'daisy'
}
```

```python
'''I know that, realistically, they're lying to themselves. In fact, the dict probably looks more like this:'''
fav_animals = {
    'parth': 'unicorn',
    'michael': 'unicorn',
    'sam': 'unicorn',
    'zheng': 'tree',
    'theo': 'unicorn',
    'alex': 'dog',
    'nick': 'daisy'
}
```

In this problem, we'll reverse the fav_animals dictionary to create a new dictionary which associates animals to a list of people for whom that animal is their favorite.

More precisely, **write a function that properly reverses the keys and values of a dictionary** - each key (originally a value) should map to a collection of values (originally keys) that mapped to it. Example in the cell:

```python
flip_dict({"CA": "US", "NY": "US", "ON": "CA"}) => {"US": ["CA", "NY"], "CA": ["ON"]}
```

<!-- grade_id: cell-423c5e21cf539125, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
def flip_dict(d):
raise NotImplementedError

### BEGIN SOLUTION

def flip_dict(d):
    out = {}
    for key, value in d.items():
        if value not in out:
            out[value] = []
        out[value].append(key)
    return out

### END SOLUTION
```

Output:
```
{'unicorn': ['parth', 'michael', 'sam', 'theo'],
 'tree': ['zheng'],
 'dog': ['alex'],
 'daisy': ['nick']}
```

<!-- grade_id: cell-f7e4b164b1cee83b, points: 0, grade: true, solution: false -->
```python
assert flip_dict({"CA": "US", "NY": "US", "ON": "CA"}) == {"US": ["CA", "NY"], "CA": ["ON"]}

### BEGIN HIDDEN TESTS

assert flip_dict({
    'parth': 'unicorn',
    'michael': 'unicorn',
    'sam': 'unicorn',
    'zheng': 'tree',
    'theo': 'unicorn',
    'alex': 'dog',
    'nick': 'daisy'}) == {'unicorn': ['parth', 'michael', 'sam', 'theo'],
                          'tree': ['zheng'],
                          'dog': ['alex'],
                          'daisy': ['nick']}

### END HIDDEN TESTS
```